---
title: Sample Exercises
description: >-
  These exercises will teach you how to write your own exercises. Notice how the
  Chapter title is written in Title Case.

--- type:VideoExercise lang:r xp:50 skills:1 key:e897462746
## A Video Exercise

Video exercises should just contain bullet points describing the things you want to talk about.

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:fb611e1b51
## A Multiple Choice Exercise

Multiple choice exercises

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:f54715376f
## A Coding Exercise

Coding exercises should form the bulk of the course.

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
